# SLU16 - Workflow: Exercise Notebook

```python
import random
import pandas as pd
import matplotlib.pyplot as plt

from sklearn import linear_model
from sklearn.pipeline import make_pipeline
from sklearn.preprocessing import MinMaxScaler
from sklearn.ensemble import RandomForestClassifier
from sklearn.metrics import accuracy_score
from sklearn.base import TransformerMixin
from sklearn.model_selection import train_test_split
from sklearn.impute import SimpleImputer

from hashlib import sha1 # just for grading purposes
import json # just for grading purposes

from utils import workflow_steps, data_analysis_steps

def _hash(obj, salt='none'):
    if type(obj) is not str:
        obj = json.dumps(obj)
    to_encode = obj + salt
    return sha1(to_encode.encode()).hexdigest()
import warnings
warnings.filterwarnings('ignore')
```

In this notebook we will be covering the following:

* Workflow
 * Step 1: Get the data
 * Step 2: Data analysis and preparation
   * 2.1 Data analysis
   * 2.2 Dealing with data problems
   * 2.3 Feature engineering
   * 2.4 Feature selection
 * Step 3: Model training
 * Step 4: Evaluate results
* Pipelines and Custom estimators
    * Pipelines 
        * Doing it "the hard way"
        * What is a pipeline
        * Setting up a pipeline
    * Custom Objects
        * Custom Transformers
        * Custom Estimators

A data science **workflow** defines the phases (or steps) in a data science project. Using a well-defined data science workflow is useful not only to you, but your teammates as it provides a simple way to clearly structure and organize a data science project. Across this specialization we've been covering the different steps in this workflow, but how well are you familiarised with them?

## Exercise 1: Workflow

### Exercise 1.1 - Overall workflow steps

What are the basic workflow steps?

You probably know them already, but we want you to really internalize them. We've given you a list of steps in `workflow_steps`, but it appears that, not only does it have to many steps, some are _probably_ wrong, as well.

Select the correct ones and reorder them!

```python
print("Workflow steps:")
for i in range(len(workflow_steps)):
    print(i+1, ': ', workflow_steps[i])
```

<!-- grade_id: cell-eb636145aa25eef4, points: 0, grade: false, solution: true -->
```python
# Exercise 1.1. Filter and sort the names of the steps in the workflow_steps list
# workflow_steps_answer = [...]

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8c9ae302d2c9c48c, points: 1, grade: true, solution: false -->
```python
assert _hash(len(workflow_steps_answer), 'salt0') == '253bd16d7f97f7d7295d404825e32985d257e11f', "your workflow size doesn't look right! Don't forget to remove steps that shouldn't be there"
assert _hash([step.lower() for step in workflow_steps_answer], 'salt0') == '701e2306da9bfde36382bdb6feb80a354916ebf4', "your workflow order doesn't look right! Some steps might be out of place"

```

## Exercise 1.2: Data Analysis and Preparation

There are way too many substeps in the **Data Analysis and Preparation** step to group them all under a single category. We've given you another list of steps: `data_analysis_steps`.

Aside from being shuffled, it should be fine but keep an eye out. You never know what to expect...

```python
print("Data Analysis and Preparation steps:")
for i in range(len(data_analysis_steps)):
    print(i+1, ': ', data_analysis_steps[i])
```

<!-- grade_id: cell-2dab17640925e0d9, points: 0, grade: false, solution: true -->
```python
# Exercise 1.2. Filter and sort the names of the steps in the data_analysis_steps list
# data_analysis_steps_answer = [...]

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6d6087c6110d2371, points: 1, grade: true, solution: false -->
```python
assert _hash(len(data_analysis_steps_answer), 'salt0') == '3aa298b947129ff226ba1e3f545f8b5ffaa6e557', "your workflow size doesn't look right! Don't forget to remove steps that shouldn't be there"
assert _hash([step.lower() for step in data_analysis_steps_answer], 'salt0') == '658ab90eff4a0cea2bfb51cc89c8db5b4121fa86', "your workflow order doesn't look right! Some steps might be out of place"

```

<img src="media/spanish_inquisition.gif" width="400" />

## Exercise 2 - Walking down the yellow (workflow) path

There is no template for solving a data science problem. The roadmap changes with every new dataset and new problem. But we do see similar steps in many different projects. Regardless, some steps are fairly common in any process. Let's go through them one by one.

### Exercise 2.1 - Objective

Every DS analysis should start with one question: **What is the problem you are trying to solve?** Clearly stating your problem is the first step to solving it and without a clear problem, you could find yourself down a data-science rabbit-hole.

For this workflow, we are going to analyze the **highest ranked movies on IMDB.com**. The objective? 

> I want to build a model to predict IMDB movie rating based on features like budget, runtime, and votes on the website.

#### Exercise 2.1.1 - Objective
Let's start by importing the dataset. 

```python
df = pd.read_csv('data/IMDB.csv', index_col=0, encoding='latin-1')
df.head()
```

<!-- grade_id: cell-7bac9e5a5b30e333, points: 0, grade: false, solution: true -->
```python
# Is the objective clear to you?
# This is just a yes or no question, no need for code here! :P 
#answer_2_1_1=False
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-dad9e0e457bf29a1, points: 0, grade: true, solution: false -->
```python
assert answer_2_1_1, "Don't make the panda sad!"
```

#### Exercise 2.1.2 - Clasifying the problem
Now that we have our data imported into Pandas and we've checked out the first few rows of our dataframe, there's a few questions we need to answer before we move on:

- *A*: Is this **supervised learning** or **unsupervised learning**? 
- *B*: Is this a **classification problem** or is it a **regression problem**? 
- *C*: Is this a **prediction problem** or an **inference problem**?

Keeping our **objective in mind** how would you classify this problem?

In our example, we are going to be using **regression** to **predict** IMDB rating from Metacritic Ratings, Budgets, Runtime, and Votes.

Save, in `answer_2_1_2`, the value from **A, B and C** that apply to our problem!

<!-- grade_id: cell-26745fdc9dca03d8, points: 0, grade: false, solution: true -->
```python
# Remove from the string what doesn't apply to our problem (including the '/')
#answer_2_1_2 = ["supervised/unsupervised learning","classification/regression problem", "prediction/inference problem"]
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-80d1e9191be2cad6, points: 1, grade: true, solution: false -->
```python
assert _hash(len(answer_2_1_2), 'salt0') == 'bb6591f296df2e5dd0324184f4348795b760f11d', "your answer doesn't have the correct size. I've asked you to pick the correct option for three questions"
assert _hash([step.lower() for step in answer_2_1_2], 'salt0') == '316a377a65564efdcfaadfec185b3e53aa41995e', "One or more of your answers is incorrect."
```

### Exercise 2.2 - Data Exploration and Data Cleaning

Back to our data! Let’s determine which variable is our target and which features we think are important.
Our target is going to be the column titled **Rating** and our features are going to be the columns titled the following: **Rating, TotalVotes, MetaCritic, Budget, Runtime, VotesUS and VotesnUS**, in this order. Remove all of the columns we don’t need for this analysis.

Save the resulting dataframe in `df_clean`. 

Remember, in this case we're telling you what columns are to be kept, but this is something that is part of your workflow process. **A good data exploration and data cleaning is a key factor in the outcome of your model!**

<!-- grade_id: cell-138aa6edcbd29190, points: 0, grade: false, solution: true -->
```python
#df_clean = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-797514f689b49b52, points: 1, grade: true, solution: false -->
```python
assert isinstance(df_clean, pd.DataFrame), "Should be a dataframe"
assert df_clean.shape == (118, 7), "The shape of the dataframe is different then expected. Have you dropped the uncessessary columns?"
assert _hash([step.lower() for step in df_clean.columns], 'salt0') == 'a721c6ef4e309977818a8875f46459655f645b70', "One or more of your column headers is incorrect."

```

### Exercise 2.3 - EDA

Exploratory data analysis (EDA) gives the data scientist an opportunity to really learn about the data they are working with. 

Throughout the EDA process, I clean the data. Data from the real world is *very messy*. As I work through the EDA process and learn about the data, I take notes on things I need to fix in order to conduct my analysis. Most times, **Data cleaning and EDA go hand in hand for me**.

The first thing I check is data types. Getting all of the values in the correct format is important. This can involve stripping characters from strings, converting integers to floats, or many other things.

```python
df_clean.dtypes

```

We can observe that the columns **Total Votes**, **Budget** and **Runtime** are not in the expected format (they're all as objects). 

Let's fix these columns one by one:

#### Exercise 2.3.1 - Total Votes
Remove the commas in the column and set its type as `int` 

```python
df_clean.TotalVotes
```

<!-- grade_id: cell-c0fece0cf6d920ba, points: 0, grade: false, solution: true -->
```python
#hit: str.replace() and as.type() are your friends
# df_clean['TotalVotes'] = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-80f21621753787af, points: 1, grade: true, solution: false -->
```python
assert all(isinstance(item, int) for item in df_clean['TotalVotes']), "TotalVotes should be saved as integers" 
assert _hash(list(df_clean['TotalVotes']), 'salt0') == '8213634d779bf710f03e6ea7b9cafc2671d5f5c4', "Something is wrong with your data conversion"
```

#### Exercise 2.3.2 - Budget
Set the column Budget as int, cleaning the weird cases as you go

```python
df_clean.Budget
```

<!-- grade_id: cell-bde0821a4f079b78, points: 0, grade: false, solution: true -->
```python
# remove cases where the budget contains the words "Opening", "Pathé" and "BBC"
#df_clean = df_clean[...]

# remove weird string cases from the column - this is a freeby for you ;)
#df_clean.Budget = df_clean.Budget.str.replace('[^\x00-\x7F]','')

# remove commas and dollar strings
# df_clean.Budget = ..

# set type as int
# df_clean.Budget = ...


# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-733c4a2a5e210740, points: 2, grade: true, solution: false -->
```python
assert all(isinstance(item, int) for item in df_clean['Budget']), "Budget should be saved as integers" 
assert _hash(list(df_clean['Budget']), 'salt0') == '36348180fcef98c4a20e296250c54dcba9150ce4', "Something is wrong with your data cleaning"
```

#### Exercise 2.3.3 - Runtime
Extract the runtime time in minutes from this column. *Don't convert this column.* 

```python
df_clean.Runtime
```

<!-- grade_id: cell-39c6b959c9e31b30, points: 0, grade: false, solution: true -->
```python
# hint: str.split() is your friend
# df_clean.Runtime = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1c14f70638eb8f1c, points: 1, grade: true, solution: false -->
```python
assert all(isinstance(item, object) for item in df_clean['Runtime']), "Runtime should be left as object" 
assert _hash(list(df_clean['Runtime']), 'salt0') == 'd17030b4117d1ea61ead4b8b3ddfdb82a0bfc0bf', "Something is wrong with your data cleaning"
```

### Exercise 2.4 - Drop missing values
Finding missing values is quite common. Just drop them from out clean dataframe.

<!-- grade_id: cell-fc4703b89bf0d85b, points: 0, grade: false, solution: true -->
```python
#df_clean = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8443173a170a44d8, points: 1, grade: true, solution: false -->
```python
assert isinstance(df_clean, pd.DataFrame), "Should be a dataframe"
assert df_clean.shape == (92, 7), "The shape of the dataframe is different then expected. Have you dropped the rows with missing data?"
```

### Exercise 2.5 - Set type of remaining columns

Finally we can set the type of **Runtime**. as int, and **VotesnUS** and **VotesUS** as float.

<!-- grade_id: cell-53bce2cc4afac670, points: 0, grade: false, solution: true -->
```python
#df_clean.Runtime = ...
#df_clean.VotesnUS = ...
#df_clean.VotesUS = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-24f27a68d3dcd195, points: 0, grade: true, solution: false -->
```python
assert all(isinstance(item, int) for item in df_clean['Runtime']), "Runtime should be saved as integers" 
assert all(isinstance(item, float) for item in df_clean['VotesnUS']), "VotesnUS should be saved as integers" 
assert all(isinstance(item, float) for item in df_clean['VotesUS']), "VotesUS should be saved as integers" 

pd.plotting.scatter_matrix(df_clean, figsize=(20,20))
plt.show()
```

Uff! That took quite some time, but now we have a clean and tidy dataframe to work with!

## Exercise 3 - Baseline modeling

As a data scientist, you will build a lot of models. You will use a variety of algorithms to perform a wide variety of tasks. You will need to use intuition and experience to decide when certain models are appropriate! 

But when costructing your baseline model, the simpler, the better! 

Let's start!

### Exercise 3.1 - Separate your target value 

Separate into `X` and `y` your features and your target. Order the columns in `X` as in exercise 2.2. 

<!-- grade_id: cell-c1c072726b679fe8, points: 0, grade: false, solution: true -->
```python
#X = ...
#y = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d02011f7f3426b07, points: 1, grade: true, solution: false -->
```python
assert _hash(list(y), 'salt0') == 'a285222c126f8617cd0179da57e9879065213a75', "Have you picked the right column as the target?"
assert X.shape == (92, 6), "The shape of the dataframe is different then expected. Have you dropped the target?"
assert _hash(sorted(X.columns), 'salt0') == '91dadf1574a4bf5ac09cbc839a67393502645794', "Have you included the right columns in X?"
```

### Exercise 3.2 - Scale your data
As we're using a linear regression, you should scale your X dataframe. Use `MinMaxScaler`

<!-- grade_id: cell-fe58c296cca28681, points: 0, grade: false, solution: true -->
```python
# corr_X = MinMaxScaler().fit_transform(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cc5e79e6d94f3521, points: 1, grade: true, solution: false -->
```python
assert _hash(sorted(corr_X[:,1]), 'salt0') == '1e067622da29527aa7f0ea65f8c60d6eb202dd3a', "have you used the correct scaler?"

```

### Exercise 3.3 - Split data

Split your dataset into test and train data, using `test_size=0.2` on your `train_test_split()`

<!-- grade_id: cell-0a59ceb29b6e7a7b, points: 0, grade: false, solution: true -->
```python
#X_train, X_test, y_train, y_test = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-eda4965289c315c6, points: 1, grade: true, solution: false -->
```python
assert (X_train.shape, X_test.shape, y_train.shape, y_test.shape) == ((73, 6), (19, 6), (73,), (19,)), "have you split the data correctly? Test size should be 0.2"
```

### Exercise 3.4 - Finally! The model!

We can finally make our predictions with our Linear Regression! 

<!-- grade_id: cell-6423f554ffdba88d, points: 0, grade: false, solution: true -->
```python
#lm = linear_model.LinearRegression()
#model = lm.fit(..., ... )
#predictions = lm.predict(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-408ce12cfa728ad1, points: 2, grade: true, solution: false -->
```python
assert predictions.size == 19
plt.scatter(y_test, predictions)
plt.xlabel('True y')
plt.ylabel('Predicted y')
plt.show()

print ('R-squared Score:', model.score(X_test, y_test)) 
```

Our model is not performing bad at all! If you want to improve it, you should make **small alterations**, **one at a time**! Keeping track of your changes is crucial to know exactly what change is helping or hurting your model! 

## Exercise 4 - Pipelines!!!!! 

We've already loaded and splitted a dataset for the following exercises. They're stored in the `X_train`, `X_test`, `y_train` and `y_test` variables.

In a perfect world, where you have all your data clean and ready-to-go, you can create your pipeline with just Scikit-learn's Transformers. However, in the real world, that's not the case, and you'll need to create custom Transformers to get the job done. Take a look at the data set, what do you see?

```python
from utils import get_dataset

new_X, new_y = get_dataset()  # preloaded dataset
new_X_train, new_X_test, new_y_train, new_y_test = train_test_split(new_X, new_y, test_size=0.33)
```

```python
# something interesting in the data?
```

While crunching your data, you probably found two issues:

There are 4 columns whose name starts with either arm or leg which are all filled with gibberish.
There are some values missing in some columns.
So, first things first, let's get rid of those columns through a **Custom Transformer**, so we can plug it in a Scikit Pipeline after.

### Exercise 4.1 - Custom Transformer

<!-- grade_id: cell-dc3fa4786483541e, points: 0, grade: false, solution: true -->
```python
# Create a pipeline step called RemoveLimbs that removes any
# column whose name starts with the string 'arm' or ´leg´

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b91664f72233c2c7, points: 3, grade: true, solution: false -->
```python
assert _hash(sorted(RemoveLimbs().fit_transform(new_X).columns), 'salt5') == '71443dfc3077d773d4c74e958dadf91dc2cc148a'
assert _hash(list(map(lambda col: col.startswith('arm') or col.startswith('leg'), RemoveLimbs().fit_transform(new_X_train).columns)), 'salt6') == 'ce45cf3759d2210f2d1315f1673b18f34e3ac711'
```

### Exercise 4.2 - Pipelines are the best!

Now that we have our Custom Transformer in place, we can design our pipeline! 

Create a pipeline with the following steps:

1. Removes limbs columns
2. Imputes missing values with the mean
3. Has a Random Forest Classifier as the last step

You may use `make_pipeline` to create your pipeline with as many steps as you want as long as the first two are the Custom Transformer you developed previously, a `SimpleImputer` as the second step, and a `RandomForestClassifier` as the last step. Save your pipeline into a variable named `pipeline`!

<!-- grade_id: cell-178a0f61ec49c4ac, points: 0, grade: false, solution: true -->
```python
# pipeline = make_pipeline(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3939f2f56f82f742, points: 3, grade: true, solution: false -->
```python
assert _hash(pipeline.steps[0][0], 'salt7') == '471b02068ac2c4f479c2e9f85f4b3dc2179bb841'
assert _hash(pipeline.steps[1][0], 'salt8') == 'ca83eaea1a7e243fa5574cfa6f52831166ee0f32'
assert _hash(pipeline.steps[-1][0], 'salt9') == '0d66ba4309ad4939673169e74f87088dcadd510b'
```

Does it work? Let's check it out on our dataset!

```python
pipeline.fit(new_X_train, new_y_train)
new_y_pred = pipeline.predict(new_X_test)
accuracy_score(new_y_test, new_y_pred)
```

It doesn't get much cleaner that this, does it?

For an extra challenge, Go back to exercises 2 and 3 and follow our workflow but with a Pipeline! For each special processing we've done, you can create a custom transformer for that column! 

Dominating pipelines and custom transformers can be a huge time saver! And there's the hackathon ahead...

**Good luck!**