# SLU19 - k-Nearest Neighbors (kNN) -- Exercises

In this notebook we will have exercises covering the following topics:

- k-Nearest Neighbors Algorithm
- A Primer on Distance
- Some considerations about kNN
- Using kNN

In the first part of the notebook you will be implementing things from scratch, so you understand what's going on under the hood. 

![numpy-function-implementation](media/numpy-function-implementation.png)

```python
# Place any important imports at the top of the notebook when possible
import hashlib
import math
import os

import numpy as np
import pandas as pd

from sklearn import datasets

import json
from hashlib import sha1 # just for grading purposes

def _hash(obj):
    if type(obj) is not str:
        obj = json.dumps(obj)
    return sha1(obj.encode()).hexdigest()
```

## Distances

To understand the distances compute here, you first need to understand what the Euclidean norm is. You have seen it in the learning notebook, and repeated across several formulas. Yes, we're talking about this:

$$|\mathbf{x}|$$ 

So how do we define it? 

Well, let's look at the Euclidean distance equation we've shown you in the learning notebook. The norm definition is "hidden" there:

$$d(\mathbf{p}, \mathbf{q}) = \sqrt{(q_1 - p_1)^2 + (q_2 - p_2)^2 + ... + (q_n - p_n)^2} = \sqrt{ \sum_{i=1}^n (q_i - p_i)^2} = |\mathbf{q} - \mathbf{p}|$$

If you focus on the right side of the equation, it should be obvious that it contains a definition of the **euclidean norm** applied to the difference between the vectors p and q:

$$ \sqrt{ \sum_{i=1}^n (q_i - p_i)^2} = |\mathbf{q} - \mathbf{p}|$$

So we can just replace `p-q` by a single vector `x` in the equation:

$$ \sqrt{ \sum_{i=1}^n (x_i)^2} = |\mathbf{x}|$$

And we get our norm definition!


### Exercise 1 - Vector norms

Start by implementing the Euclidean norm definition you explored above:

$$|\mathbf{x}| = \sqrt{ \sum_{i=1}^n (x_i)^2} = \sqrt{(x_1)^2 + (x_2)^2 + ... + (x_N)^2}$$



<!-- grade_id: cell-3590731766111833, points: 0, grade: false, solution: true -->
```python
def euclidean_norm(x):
    """
    Return the euclidean norm of a vector
        
    Parameters
    ----------
    x: numpy array with shape (N,)
    
    Returns
    ----------
    norm: float
    """

    # YOUR CODE HERE
    raise NotImplementedError()
    return norm

```

<!-- grade_id: cell-ff9a506bc3c80751, points: 1, grade: true, solution: false -->
```python

np.testing.assert_almost_equal(euclidean_norm(np.array([1, 2, 4])), 4.5825, 2)
np.testing.assert_almost_equal(euclidean_norm(np.array([-1, 0, 4])), 4.1231, 2)
np.testing.assert_almost_equal(euclidean_norm(np.array([1])), 1.0, 2)
np.testing.assert_almost_equal(euclidean_norm(np.array([-1])), 1.0, 2)
np.testing.assert_almost_equal(euclidean_norm(np.array([0, 0])), 0.0, 2)
np.testing.assert_almost_equal(euclidean_norm(np.array([0, 1, 2, 3, 4])), 5.4772, 2)
np.testing.assert_almost_equal(euclidean_norm(np.array([0, -1, -2, -3, -4])), 5.4772, 2)

```

### Exercise 2 - Distances

Now that we have the norm, we'll apply it to find our distances. 

Define a function called `distance_function`. 

This function should receive two arguments, `a` and `b`, both numpy arrays with shape `(N,)`, where `N` is the number of dimensions of the inputs `a` and `b`, and calculate the distance between them. Additionally, it receives a keyword argument `distance_type`, which tells you which distance to use. 


The argument can have one of three values, which will define how to compute the distance (we could have picked different formulations, but we chose these for you to use the norm function created before):

* `euclidean`

$$d_{euclidean} = |\mathbf{b} - \mathbf{a}|$$


* `dot`

$$d_{dot} = u_1v_1 + u_2v_2 + ... + u_nv_n$$

* `cosine`

$$cosine(\mathbf{a}, \mathbf{b}) = 1 -  \frac{\mathbf{a} \; . \mathbf{b}}{|\mathbf{a}| \; |\mathbf{b}|}$$



<!-- grade_id: cell-6e62f4a661edbaf0, points: 0, grade: false, solution: true -->
```python
# implement a function called euclidean_distance

def distance_function(a, b, distance_type="euclidean"):
    """
    Return the distance between two vectors, computed using one of 
        `euclidean`, `dot_product` or `cosine`. 
        
    Return `None` if:
     - distance type is not any of the supported types
     - if the shape of `a` and `b` do not match

    Parameters
    ----------
    a: numpy array with shape (N,)
    b: numpy array with shape (N,)
    distance_type: str - enumerate, can be one of `euclidean`, `dot`
        or `cosine`
    
    Returns
    ----------
    distance: float
    """
    
    # 1. Check shape consistency
    # YOUR CODE HERE
    raise NotImplementedError()

    # 2. Compute distance
    # Hint:
    # 
    # if distance_type == "euclidean":
    #     distance = ...
    # elif ...
    #      ... 
    # elif ...
    #      ... 
    # else:...
    #      ...

    # YOUR CODE HERE
    raise NotImplementedError()
    
    return distance
```

<!-- grade_id: cell-79f8353bbd6472e8, points: 4, grade: true, solution: false -->
```python
# Test Euclidean Distance

np.testing.assert_almost_equal(
    distance_function(np.array([1, 2, 4]), np.array([-1, 0, 4]), distance_type="euclidean"), 2.8284, 2)

np.testing.assert_almost_equal(
    distance_function(np.array([1]), np.array([-1]), distance_type="euclidean"), 2.0000, 2)

np.testing.assert_almost_equal(
    distance_function(np.array([0, 0]), np.array([2, 3]), distance_type="euclidean"), 3.6055, 2)

np.testing.assert_almost_equal(
    distance_function(np.array([0, 1, 2, 3, 4]), np.array([0, -1, -2, -3, -4]), distance_type="euclidean"), 10.9544, 2)


# Test Dot product

np.testing.assert_almost_equal(
    distance_function(np.array([1, 2, 4]), np.array([-1, 0, 4]), distance_type="dot"), 15.0, 2)

np.testing.assert_almost_equal(
    distance_function(np.array([1]), np.array([-1]), distance_type="dot"), -1.0, 2)

np.testing.assert_almost_equal(
    distance_function(np.array([0, 0]), np.array([2, 3]), distance_type="dot"), 0.0, 2)

np.testing.assert_almost_equal(
    distance_function(np.array([0, 1, 2, 3, 4]), np.array([0, -1, -2, -3, -4]), distance_type="dot"), -30.0, 2)


# Test Cosine distance

np.testing.assert_almost_equal(
    distance_function(np.array([1, 2, 4]), np.array([-1, 0, 4]), distance_type="cosine"), 0.2061, 2)

np.testing.assert_almost_equal(
    distance_function(np.array([1]), np.array([-1]), distance_type="cosine"), 2.0, 2)

np.testing.assert_almost_equal(
    distance_function(np.array([0, 1]), np.array([2, 3]), distance_type="cosine"), 0.1679, 2)

np.testing.assert_almost_equal(
    distance_function(np.array([0, 1, 2, 3, 4]), np.array([0, -1, -2, -3, -4]), distance_type="cosine"), 2.0, 2)


# Test cases where distance can't be computed

assert distance_function(np.array([1, 2]), np.array([-1, 0, 4]), distance_type="euclidean") is None
assert distance_function(np.array([1, 2]), np.array([-1, 0, 4]), distance_type="dot_product") is None
assert distance_function(np.array([1, 2]), np.array([-1, 0, 4]), distance_type="cosine") is None
assert distance_function(np.array([1, 2, 3]), np.array([-1, 0, 4]), distance_type="no_distance") is None

```

You probably know about some nicer functions to implement what you have done here:
* [numpy.linalg.norm](https://docs.scipy.org/doc/numpy/reference/generated/numpy.linalg.norm.html)
* [numpy.dot](https://docs.scipy.org/doc/numpy/reference/generated/numpy.dot.html)
* [scipy.distance.cosine](https://docs.scipy.org/doc/scipy-0.14.0/reference/generated/scipy.spatial.distance.cosine.html)

But we want you to really try to implement these by yourself and understand what is happening.


## Implementing the kNN algorithm

Now that we have all of our distances, we'll implement the kNN algorithm.

And we'll do it by hand! Let's do this!

![lets_do_this](media/lets_do_this.gif)

### Exercise 3 - Finding the closest neighbors 

The first step is to find the nearest data points. For that purpose, implement a function called `find_nearest_neighbours`, that:

* receives four arguments:
    * x, which is a numpy array with shape (d,)
    * dataset, which is a numpy array with shape (N, d), where N is the dataset size
    * distance_type, which can be 'euclidean', 'cosine', 'dot'
    * k, which is the number of nearest neighbors that we want to consider
* Iterates through the dataset and computes the distance from point x to each dataset point, using the function you build before (`distance_function`)
* gets the indexes of the k smallest distances (in ascending order)
* returns a numpy array of shape (k,) with those indexes

Hint: check [numpy.argsort](https://docs.scipy.org/doc/numpy/reference/generated/numpy.argsort.html).

<!-- grade_id: cell-579a133d5f7272c2, points: 0, grade: false, solution: true -->
```python
def find_nearest_neighbours(x, dataset, distance_type="euclidean", k=5):
    """
    Finds the k nearest neighbors by getting a distance between 
    a point and all the other points in a dataset and sorting 
    them, retrieving the indices of the closest points
    
    Parameters
    ----------
    x: numpy array with shape (d,)
    dataset: numpy array with shape (N, d)
    distance_type: str - enumerate, can be one of `euclidean`, `dot_product`
        or `cosine`
    k: int, the number of nearest neighbors we want to consider

    Returns
    ----------
    indexes: numpy array with shape (k,)
    """

    # 1. Compute the distance from x to all points in dataset
    # YOUR CODE HERE
    raise NotImplementedError()
      
    # 2. Get the indices of the closest data points
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return indexes

```

<!-- grade_id: cell-601006c4a24f1f22, points: 3, grade: true, solution: false -->
```python
dataset = datasets.load_iris().data
x = np.array([4.9, 3.0, 6.1, 2.2])

knn_1 = find_nearest_neighbours(x, dataset, 'euclidean', 3)
assert knn_1.shape == (3,)
assert _hash(knn_1.tolist()) == '408b0d26e3ebc7c7998cd6ae42215e0f49c2f867'

knn_2 = find_nearest_neighbours(x, dataset, 'euclidean', 10)
assert knn_2.shape == (10,)
assert _hash(knn_2.tolist()) == '637f612aec53627ec340ab031592ce29e119ae71'

knn_3 = find_nearest_neighbours(x, dataset, 'dot', 3)
assert knn_3.shape == (3,)
assert _hash(knn_3.tolist()) == '45b13c85dbee9e89703048b3b455afc6bee87b8e'

knn_4 = find_nearest_neighbours(x, dataset, 'dot', 10)
assert knn_4.shape == (10,)
assert _hash(knn_4.tolist()) == 'f4399de12191b8025964f37dc394a6efc1dd937f'

knn_5 = find_nearest_neighbours(x, dataset, 'cosine', 3)
assert knn_5.shape == (3,)
assert _hash(knn_5.tolist()) == '6dcf71640710b0aa90cae6e856220c8919a167ab'

knn_6 = find_nearest_neighbours(x, dataset, 'cosine', 10)
assert knn_6.shape == (10,)
assert _hash(knn_6.tolist()) == '98a8ccaafdfaab3aba3a122fd22982cb820c6370'

```

### Exercise 4 - Classifying from nearest neighbours

Now that we have a function that gets the indexes of the k nearest neighbors, we need to get the values of those neighbors, so that afterwards we can predict the label for our point. For the classification problem, this 
means getting all the labels (the values from the neighbours) and **returning the most common label**.

In this exercise, you'll implement a function called `get_knn_class`, where you'll do just that:

* receives two arguments:
    * y, which is a numpy array with the targets from dataset
    * neighbor_indexes, which are the indexes of the k nearest neighbors (like the output of the last function)
* gets the values from y using the indexes from neighbor_indexes
* checks the most frequent label and returns it


<!-- grade_id: cell-f9d4962cb0206df9, points: 0, grade: false, solution: true -->
```python
# implement a function called get_neighbors_labels

def get_knn_class(y, neighbor_indexes):
    """
    Get the label values from the k nearest neighbors and 
    return the most frequent one, i.e., the actual class
    that knn classification yields.
    
    Parameters
    ----------
    y: numpy array with shape (N,) - labels
    neighbor_indexes: numpy array with shape (k,) - indexes of neighbors
    
    Returns
    ----------
    knn_label: int
    """
    

    # 1. Get the labels from the observed neighbors
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # 2. Get the most common label
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return knn_label

```

<!-- grade_id: cell-10829453a40bdac6, points: 2, grade: true, solution: false -->
```python
np.random.seed(42) 

# Test case 1
answer = get_knn_class(np.random.rand(150), np.random.randint(0, 3, 3))
assert isinstance(answer, int)
assert answer == 0, answer

# Test case 2
answer = get_knn_class(np.random.rand(10), np.random.randint(1, 5, 7))
assert isinstance(answer, int)
assert answer == 0, answer

np.random.seed(42) 

```

### Exercise 5 - Classification with KNN (Putting everything together)

Finally we can put everything together and implement the knn classifier!

In this exercise, you'll implement a function called `knn_classifier`, that:

* receives five arguments:
    * x, which is a numpy array with shape (d,)
    * dataset, which is a numpy array with shape (N, d), where N is the dataset size
    * targets, which is a numpy array with shape (N,), that has the targets for each of the points in the dataset
    * k, which is the number of nearest neighbors our knn algorithm will consider
    * distance_function, which can be 'euclidean', 'cosine', 'dot'
* uses the functions that we implemented above in order to implement a knn_classifier!


<!-- grade_id: cell-52fb6a8a51b6ff92, points: 0, grade: false, solution: true -->
```python
def knn_classifier(x, dataset, targets, k, distance_function):
    """
    Predicts the label for a single point, given a dataset, a value for k and a distance function
    
    Parameters
    ----------
    x: numpy array with shape (d,)
    dataset: numpy array with shape (N, d)
    targets: numpy array with shape (N,)
    k: int
    distance_function: string
    
    Returns
    ----------
    label: int
    """

    # YOUR CODE HERE
    raise NotImplementedError()
    
    return label

```

<!-- grade_id: cell-089aad803d5ae57f, points: 2, grade: true, solution: false -->
```python
dataset = datasets.load_iris().data
targets = datasets.load_iris().target
x = np.array([4.9, 3.0, 6.1, 2.2])

tests = [
    {
        'input': [x, dataset, targets, 3, 'euclidean'],
        'expected_value': 2
    },
    {
        'input': [x, dataset, targets, 5, 'dot'],
        'expected_value': 0
    },
    {
        'input': [x, dataset, targets, 1, 'cosine'],
        'expected_value': 2
    }
]

for test in tests:
    pred_label = knn_classifier(*test['input'])
    assert isinstance(pred_label, int), "The function should return an integer!"
    assert pred_label == test['expected_value'], "The returned int has the wrong value!"
```

Great job! You now have a working KNN classifier!

![its-alive](media/its-alive.png)


Now that we've implemented a knn classifier, let's go a bit further and implement a knn regressor!


### Exercise 6 - Regression with KNN


As we explained in the learning notebook, the main difference between a knn classifier and a knn regressor is the way we choose the predicted label from the labels of the nearest neighbors. So we can reuse the first step of retrieving the neighbours.

For the classifier case we used a majority vote. In the regressor case, we want to use the average value of the neighbors' labels.

In this exercise, start by implementing a function called `get_knn_value`, that:

* receives two arguments:
    * y, which is a numpy array with the targets from dataset
    * neighbor_indexes, which are the indexes of the k nearest neighbors (like the output of the last function)
* gets the values from y using the indexes from neighbor_indexes
* returns the average of the nearest neighbors' labels



<!-- grade_id: cell-2ee17f126f1a73d7, points: 0, grade: false, solution: true -->
```python
def get_knn_value(y, neighbor_indexes):
    """
    Get the label values from the k nearest neighbors and 
    return the average value, i.e., the actual output of our 
    knn regressor
    
    Parameters
    ----------
    y: numpy array with shape (N,) - labels
    neighbor_indexes: numpy array with shape (k,) - indexes of neighbors
    
    Returns
    ----------
    knn_label: float
    """

    # 1. Get the labels from the observed neighbors
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # 2. Get the average label value
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return knn_label

```

<!-- grade_id: cell-dea39d8bb2bae597, points: 2, grade: true, solution: false -->
```python
np.random.seed(42) 

# Test case 1
answer = get_knn_value(np.random.rand(150), np.random.randint(0, 3, 3))
assert isinstance(answer, float)
np.testing.assert_almost_equal(answer, 0.4937, 2)

# Test case 2
answer = get_knn_value(np.random.rand(10), np.random.randint(1, 5, 7))
assert isinstance(answer, float)
np.testing.assert_almost_equal(answer, 0.5192, 2)

```

### Exercise 7

And we're ready to implement the knn regressor! Keep up the good work, we're almost there!

![almost_there](media/almost_there.gif)

In this exercise, you'll implement a function called `knn_regressor`, that:

* receives five arguments:
    * x, which is a numpy array with shape (d,)
    * dataset, which is a numpy array with shape (N, d), where N is the dataset size, and d is the number of dimensions that the points in the dataset have
    * targets, which is a numpy array with shape (N,), that has the targets for each of the points in the dataset
    * k, which is the number of nearest neighbors our knn algorithm will consider
    * distance_function, which can be 'euclidean', 'cosine', 'dot'
* uses the functions that we implemented above in order to implement a knn_regressor!

<!-- grade_id: cell-5cb72f729002e93f, points: 0, grade: false, solution: true -->
```python
# implement a function called knn_classifier

def knn_regressor(x, dataset, targets, k, distance_function):
    """
    Predicts the label for a single point, given a dataset, a value for k and a distance function
    
    Parameters
    ----------
    x: numpy array with shape (d,)
    dataset: numpy array with shape (N, d)
    targets: numpy array with shape (N,)
    k: int
    distance_function: string
    
    Returns
    ----------
    label: float
    """

    # YOUR CODE HERE
    raise NotImplementedError()

    return label

```

<!-- grade_id: cell-90a7594798c1c9fe, points: 1, grade: true, solution: false -->
```python
np.random.seed(42)
dataset = datasets.load_diabetes().data
targets = datasets.load_diabetes().target
x = np.random.rand(10)

prediction = knn_regressor(x, dataset, targets, 3, 'euclidean')
assert isinstance(prediction, float)
assert math.isclose(prediction, 265.666, rel_tol=1e-04)

prediction = knn_regressor(x, dataset, targets, 5, 'dot')
assert isinstance(prediction, float)
assert math.isclose(prediction, 92.8, rel_tol=1e-04)

prediction = knn_regressor(x, dataset, targets, 1, 'cosine')
assert isinstance(prediction, float)
assert math.isclose(prediction, 264.0, rel_tol=1e-04)
```

**Well done!!!**

![we_did_it](media/we_did_it.gif)

Finally let's wrap this up with a couple of exercises on how to use scikit's knn models.

## Using scikit's knn models

### Exercise 8

Use a `KNeighborsClassifier` to create predictions for the [brest cancer dataset](https://scikit-learn.org/stable/datasets/index.html#breast-cancer-dataset).

Please read the link above in order to understand the task we're solving 

Follow the instructions in the comments in the exercise cell.

```python
import numpy as np
import pandas as pd
import hashlib
import json

from scipy.spatial.distance import cosine
from sklearn.model_selection import train_test_split
from sklearn.neighbors import KNeighborsClassifier, KNeighborsRegressor
from sklearn.metrics import accuracy_score
```

<!-- grade_id: cell-08dc6b8c978617eb, points: 0, grade: false, solution: true -->
```python
# We start by importing the dataset
data = datasets.load_breast_cancer()

# Now do a train test split, using the train_test_split function from scikit
# Use a test_size of 0.25 and a random_state of 42
# X_train, X_test, y_train, y_test = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-85dd231d39ba1054, points: 1, grade: true, solution: false -->
```python
tests = [
    {
        'dataset_type': 'X_train',
        'dataset': X_train,
        'shape_hash': '31ffabcaf98971831a5f8ad05ba70049a86bd60bda0a971ca9691388f9f72f8b'
    },
    {
        'dataset_type': 'X_test',
        'dataset': X_test,
        'shape_hash': '747c580b9756b4741bfbe812b8ca9fd8d047a5d6f9e3ebe53d4d15117f42ec2a'
    },
    {
        'dataset_type': 'y_train',
        'dataset': y_train,
        'shape_hash': '23a4f6ee909897142105a6577ac39ff86c353b8ad0ded0bece87829bb1953a58'
    },
    {
        'dataset_type': 'y_test',
        'dataset': y_test,
        'shape_hash': '40957487610d92ca4dd2d37ec155c40d20091a504bf65270a3cd28e6863ef633'
    },
]

for test in tests:
    shape_hash = hashlib.sha256(json.dumps(test['dataset'].shape).encode()).hexdigest()

    assert isinstance(test['dataset'], np.ndarray), f"{test['dataset_type']} should be a numpy array!"
    assert shape_hash == test['shape_hash'], "The returned numpy array has the wrong shape!"
```

<!-- grade_id: cell-bf5bdec884565df8, points: 0, grade: false, solution: true -->
```python
# Now instantiate a kNN Classifier with k=3, that uses the euclidean distance as distance function
# In scikit, the euclidean distance is the default one and goes by the name of 'minkowski'
# which is in fact a generalisation of the euclidean distance
# clf = ...
# YOUR CODE HERE
raise NotImplementedError()


# Get predictions for the test dataset
# y_pred = ...
# YOUR CODE HERE
raise NotImplementedError()

# Measure the accuracy of your solution using scikit's accuracy_score function
# accuracy = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-04b5eaa6dd765f64, points: 1, grade: true, solution: false -->
```python
assert isinstance(clf, KNeighborsClassifier)
assert clf.n_neighbors == 3
assert clf.metric == 'minkowski'

assert isinstance(y_pred, np.ndarray)
assert y_pred.shape == (143,)

assert isinstance(accuracy, float)
assert math.isclose(accuracy, 0.930069, rel_tol=1e-04)
```

## Exercise 9

Now we want to see the difference if we use the cosine distance instead of the euclidean distance.

Go through the same steps as the previous exercise, but use the cosine distance as the distance metric in the knn classifier.

<!-- grade_id: cell-2d3c89f8bdfa66f9, points: 0, grade: false, solution: true -->
```python
# Instantiate a kNN Classifier with k=3, that uses the cosine distance as distance function
# clf = ...
# YOUR CODE HERE
raise NotImplementedError()


# Get predictions for the test dataset
# y_pred = ...
# YOUR CODE HERE
raise NotImplementedError()

# Measure the accuracy of your solution using scikit's accuracy_score function
# accuracy = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-27bbc70efc45d988, points: 1, grade: true, solution: false -->
```python
assert isinstance(clf, KNeighborsClassifier)
assert clf.n_neighbors == 3
assert clf.metric == cosine

assert isinstance(y_pred, np.ndarray)
assert y_pred.shape == (143,)

assert isinstance(accuracy, float)
assert math.isclose(accuracy, 0.93706, rel_tol=1e-04)
```

## Exercise 10

And the last exercise. 

Try different combinations of n_neighbors and metric and choose the option with the highest accuracy:

1. n_neighbors = 7, metric = 'minkowski'
2. n_neighbors = 9, metric = 'cosine'
3. n_neighbors = 11, metric = 'minkowski'
4. n_neighbors = 11, metric = 'cosine'

Write the answer to a variable called best_parameters as an integer (1, 2, 3 or 4)

<!-- grade_id: cell-121669e3e8e57b20, points: 0, grade: false, solution: true -->
```python
# Find the best combination of n_neighbors and metric
# best_parameters = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-75ff4cbbd03e8e9c, points: 2, grade: true, solution: false -->
```python
# Test
assert isinstance(best_parameters, int)
assert hashlib.sha256(bytes(best_parameters)).hexdigest() == '709e80c88487a2411e1ee4dfb9f22a861492d20c4765150c0c794abd70f8147c'

```

And we're done! Nice job ;)

![were_done](media/were_done.gif)