# SLU02 - Subsetting data in pandas: Exercise notebook

In this notebook you'll practice the concepts you've seen on the Learning and example notebooks:

    - Setting pandas Dataframe index
    - Selecting columns with brackets notation
    - Selecting columns with dot notation
    - Selecting rows with loc 
    - Selecting rows with iloc
    - Multi-axis indexing (as an alternative to Chain indexing) 
    - Masks
    - Where
    - Subsetting on conditions
    - Removing and Adding columns

**Note**: For these exercices we will be using a kaggle competition dataset containing the answers of several kaggle competitors to multiple choice questions. In each exercise, you'll be asked to implement a function, which you can check against the assertions provided.

Let's dive right in.

```python
# Imports needed
import pandas as pd
import numpy as np
import math
import hashlib
from utils import draw_base_puzzle, draw_final_puzzle

from IPython.core.display import HTML

# Set max rows for pandas display
pd.options.display.max_rows = 10

# Set matplotlib figures to show inline in the notebook and to be centered by default
%matplotlib inline
HTML("""
<style>
.output_png {
    display: table-cell;
    text-align: center;
    vertical-align: middle;
}
</style>
""")
```

## Kaggle Competition


As an aspiring data scientist you are eager to apply your new skills and you decide to participate in a kaggle competition. However, this competition has a twist: you must prove to have the minimum skills to enter it by completing a first data science based challenge. Easy, right?


<img src="images/kaggle_in_kaggle.png" alt="kaggle_in_kaggle" width="40%"/>


So you dive right into it. The assignment is the following: you must successfully complete a cross word puzzle where a set of hints require you to perform simple exploration tasks over a provided dataset. After completing all the words, you'll extract a secret keyphrase (marked in blue) that will unlock the competition for you.

Load the puzzle below 

```python
draw_base_puzzle()
```

As you can seee, it's a pretty simple one. You will be given 10 clues to fill each of the columns and extract the horizontal words in blue. 

But before that, start by loading the dataset you will be using to explore and complete the puzzle:

```python
# Read kaggle multiple choice dataset with new column as index
kaggle_dataset = pd.read_csv('data/multiple_choice_2017.csv')

# Show first 5 lines
kaggle_dataset.head(5)

```

Now let's dive into the clues! 


### Clue 1 - First value for `EmployerIndustry` when sorting the dataset by this column 

The first clue asks for the first value (using its natural order) of the column `EmployerIndustry` after we sort 
the dataframe by the same column. 

To solve this, start by implementing a function to change the index and sort it.

**Hint**: Keep in mind that we don't want to discard the original index (which automatically got the name `index`) as it may be useful on the long run

<!-- grade_id: select_rows, points: 0, grade: false, solution: true -->
```python
def change_and_sort_index(df, column):
    """ 
    Change dataframe index to desired column and sort, avoiding 
     repeated columns
    
    Args:
        df (pd.DataFrame): the input DataFrame
        column: column name to use as index

    Returns:p
        (pd.DataFrame): resulting Dataframe

    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: select_rows_tests, points: 2, grade: true, solution: false -->
```python
# This cell is what will test your code, please ignore it!

kaggle_dataset = pd.read_csv('data/multiple_choice_2017.csv')
new_dataset = change_and_sort_index(kaggle_dataset, "CurrentJobTitleSelect")

assert isinstance(new_dataset, pd.DataFrame)
assert new_dataset.shape[0] == kaggle_dataset.shape[0]
assert new_dataset.shape[1] == kaggle_dataset.shape[1]
assert new_dataset.index.name == "CurrentJobTitleSelect"
assert new_dataset.Country[0] == "United States"
assert 'index' in list(new_dataset.columns)

```

You'll now use the function you built to store the first clue

```python
clue_dataset = change_and_sort_index(kaggle_dataset, "EmployerIndustry")

first_sorted_employer_industry = clue_dataset.index[0]

FIRST_CLUE = first_sorted_employer_industry
```

### Clue 2 - Top country of origin for people with "Programmer" job title  


You now want to find the most common country (or top country as the clue puts it) for the subset of our data where
the `CurrentJobTitleSelect` equals `Programmer`. To do this, you decide to reuse your previous function to change the index and create an extra function to select rows by index values.

Implement below this second function and check it passes all assertions given.

<!-- grade_id: select_rows_by_country, points: 0, grade: false, solution: true -->
```python
def select_rows_from_index(df, ids):
    """ 
    Select the desired rows given the index value
    
    Args:
        df (pd.DataFrame): the input DataFrame
        ids: list with the desired values to retrieve

    Returns:
        (pd.DataFrame): subset Dataframe

    """    
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: select_rows_by_country_tests, points: 2, grade: true, solution: false -->
```python
# This cell is what will test your code, please ignore it!

kaggle_dataset = pd.read_csv('data/multiple_choice_2017.csv')
indexed_dataset = change_and_sort_index(kaggle_dataset, "CurrentJobTitleSelect")

desired_jobs = ['Computer Scientist', 'Programmer']
filtered_dataset = select_rows_from_index(indexed_dataset, desired_jobs)

assert isinstance(filtered_dataset, pd.DataFrame)
assert filtered_dataset.shape[0] == 797
assert filtered_dataset.shape[1] == kaggle_dataset.shape[1]

assert list(filtered_dataset.index.unique()) == desired_jobs

```

Combining the two you now get the second clue

```python
clue_dataset = select_rows_from_index(indexed_dataset, ['Programmer'])

# When using `value_counts` the index becomes the column values and the counts are ordered from higheest to lowest
top_programmer_country = clue_dataset.Country.value_counts().index[0] 

SECOND_CLUE = top_programmer_country
```

### Clue 3 - The top remote work status in the examples 1000 to 1200


The next clue asks for a specific example range. For this purpose, build a function that provides indices - `start` and `end` - and retrieves the slice of rows from `start` to `end`. Implement it below


<!-- grade_id: cell-8e5ea2a645f600a8, points: 0, grade: false, solution: true -->
```python
def get_slice(df, start, end):
    """ 
    Get a range of rows from the provided dataset 
    
    Args:
        df (pd.DataFrame): the input DataFrame
        start: start position for range
        end: end position for range

    Returns:
        (pd.DataFrame): subseted Dataframe

    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-710bb8203e0a0d98, points: 2, grade: true, solution: false -->
```python
# This cell is what will test your code, please ignore it!

kaggle_dataset = pd.read_csv('data/multiple_choice_2017.csv')

top_rows_1 = get_slice(kaggle_dataset, 231, 638)

assert top_rows_1.shape[0] == 407
assert top_rows_1.shape[1] == kaggle_dataset.shape[1]
assert top_rows_1.RemoteWork.values[1] == 'Most of the time'

top_rows_2 = get_slice(kaggle_dataset, 1231, 8413)

assert top_rows_2.shape[0] == 7182
assert top_rows_2.shape[1] == kaggle_dataset.shape[1]
assert top_rows_2.RemoteWork.values[4] == 'Always'

```

Use the function to get the third clue:

```python
clue_dataset = get_slice(kaggle_dataset, 300, 1200)

# When using `value_counts` the index becomes the column values and the counts are ordered from higheest to lowest
common_remote_work = clue_dataset.RemoteWork.value_counts().index[0] 

THIRD_CLUE = common_remote_work
```

### Clue 4 - The remote work status of example 4


To solve this clue, combine what you've learned above and create a function that allows you to index both the rows and the columns of a dataframe. Implement it under the function below:


<!-- grade_id: exercise_4, points: 0, grade: false, solution: true -->
```python
def dedicated_subset(df, rows, columns):
    """ 
    Select columns and rows from dataframe.
    
    Args:
        df (pd.DataFrame): the input DataFrame
        rows: list of rows to fetch
        columns: list of columns to fetch

    Returns:
        (pd.DataFrame): subsetted df

    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: exercise_4_correct, points: 2, grade: true, solution: false -->
```python
# This cell is what will test your code, please ignore it!

kaggle_dataset = pd.read_csv('data/multiple_choice_2017.csv')

kaggle_subset_single = dedicated_subset(kaggle_dataset, [1412], ['Country'])

assert kaggle_subset_single.shape[0] == 1
assert kaggle_subset_single.shape[1] == 1
assert kaggle_subset_single.Country.values[0] == 'Singapore'

kaggle_subset_multiple = dedicated_subset(kaggle_dataset, [11, 1111], ['RemoteWork', 'Age', 'Country'])

assert kaggle_subset_multiple.shape[0] == 2
assert kaggle_subset_multiple.shape[1] == 3
assert kaggle_subset_multiple.Age.values[1] == 36

```

Use the function to get the forth clue:

```python
clue_dataset = dedicated_subset(kaggle_dataset, [4], ["RemoteWork"])
clue_dataset.values[0][0]

FORTH_CLUE = clue_dataset.values[0][0]
```

Now that you have found a couple of clues, let's check how the board looks. You should be able to see something shapin up.


```python
draw_final_puzzle([FIRST_CLUE, SECOND_CLUE, THIRD_CLUE, FORTH_CLUE, "", "", "", "", "", ""])

```

Uhh exciting. Dive into the next clues to fill out the rest of the puzzle!


<img src="images/excited.gif" alt="excited" width="40%"/>



## Clues 5 - 8

These next clues ask for more complex datasets.

We will call these subset A and B. 

In particular, the clues define dataset A as follows:

* we get only coders (`CodeWriter`)
* these people are in between and including 20 and 30 years
* they are not from the United States
* they believe Univesity is important ('Very important', 'Important', 'Somewhat important')
* Contains columns 'EmploymentStatus', 'TitleFit' and 'Country'

Implement a function that obtains this dataset.


**Hint**: Look into the function `.isin` described in the learning notebook and how it can be used to check a value against a list

**Note**: Assume the df provided into these functions will be the original dataset, and as such will have the same columns without any change

<!-- grade_id: cell-a8cc62331e776b85, points: 0, grade: false, solution: true -->
```python
def kaggle_dataset_A(df):
    """ 
    Show kaggle competitors that fit the following parameters:
    
      - only coders (`CodeWriter`)
      - in between and including 20 and 30 years
      - not from the United States
      - that believe Univesity is important ('Very important', 'Important', 'Somewhat important')
      
    Return only the employment status, title fit and country
    
    Args:
        df (pd.DataFrame): the input DataFrame

    Returns:
        (pd.DataFrame): subsetted df

    """

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: exercice_5_correct, points: 4, grade: true, solution: false -->
```python
# This cell is what will test your code, please ignore it!

kaggle_dataset = pd.read_csv('data/multiple_choice_2017.csv')
kaggle_subset_A = kaggle_dataset_A(kaggle_dataset)

assert kaggle_subset_A.shape[0] == 2268
assert kaggle_subset_A.shape[1] == 3
assert kaggle_subset_A.columns.tolist() == ['EmploymentStatus', 'TitleFit', 'Country']

assert kaggle_subset_A.iloc[0].TitleFit == 'Fine'
assert kaggle_subset_A.iloc[-1].TitleFit == 'Fine'

```

Additionally, the clues define dataset B as follows:

* not from Academia (`EmployerIndustry`)
* older than 30 years
* that are from any country
* that believe University is not important (not included in 'Very important', 'Important', 'Somewhat important')
* Contains columns 'EmploymentStatus', 'TitleFit', 'CurrentJobTitleSelect' and 'Country'

Implement a function that obtains this dataset.


**Hint**: Look into the function `.isin` described in the learning notebook and how it can be used to check a value against a list

<!-- grade_id: cell-b71b61b40b8dca44, points: 0, grade: false, solution: true -->
```python
def kaggle_dataset_B(df):
    """ 
    Show kaggle competitors that fit the following parameters:
    
      - not from Academia (`EmployerIndustry` is not `Academic`)
      - older than 30 years
      - that are from any country
      - that believe University is not important (not included in 'Very important', 'Important', 'Somewhat important')
      
    Return only the employment status, title fit, current job title and country
    
    Args:
        df (pd.DataFrame): the input DataFrame

    Returns:
        (pd.DataFrame): subsetted df

    """

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-edb899ebe53259f3, points: 4, grade: true, solution: false -->
```python
# This cell is what will test your code, please ignore it!

kaggle_dataset = pd.read_csv('data/multiple_choice_2017.csv')
kaggle_subset_B = kaggle_dataset_B(kaggle_dataset)

assert kaggle_subset_B.shape[0] == 3561
assert kaggle_subset_B.shape[1] == 4
assert kaggle_subset_B.columns.tolist() == ['EmploymentStatus', 'TitleFit', 'CurrentJobTitleSelect', 'Country']

assert kaggle_subset_B.iloc[0].CurrentJobTitleSelect == 'Business Analyst'
assert kaggle_subset_B.iloc[-1].CurrentJobTitleSelect == 'DBA/Database Engineer'

```

Now look into the clues and retrieve the correct values to fill the puzzle:

* **Clue 5 - The country of the 14th element of subset B**
* **Clue 6 - The country of the 142th element of subset A**
* **Clue 7 - The country of the 9th element of subset B**
* **Clue 8 - The title fit of the 7th element of subset A**


```python
FIFTH_CLUE = kaggle_subset_B.Country.values[14]
SIXTH_CLUE = kaggle_subset_A.Country.values[142]
SEVENTH_CLUE = kaggle_subset_B.EmploymentStatus.values[9]
EIGTH_CLUE = kaggle_subset_A.TitleFit.values[7]

```

## Clues 9-10

These next clues once again require that you filter your dataset into different subsets.

We will call these subset C and D. In particular:

* dataset C should not contain people with job title `Programmer`
* dataset D should contain only people with job title `Computer Scientist` 

As you probably realize, you could use functions similar to the ones before to retrieve this dataset. However, we want you to use what you have learned regarding hiding data. 

Implement functions that hide the non-desired data but keep the dataframe shape.


<!-- grade_id: cell-dd32ace8f3921366, points: 0, grade: false, solution: true -->
```python
def kaggle_dataset_C(df):
    """ 
    Hide all Programmers in dataframe
    
    Args:
        df (pd.DataFrame): the input DataFrame

    Returns:
        (pd.DataFrame): output DataFrame

    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d87eef42c4baa2a2, points: 2, grade: true, solution: false -->
```python
# This cell is what will test your code, please ignore it!

kaggle_dataset = pd.read_csv('data/multiple_choice_2017.csv')
kaggle_subset_C = kaggle_dataset_C(kaggle_dataset)

assert isinstance(kaggle_subset_C, pd.DataFrame)
assert kaggle_subset_C.shape[0] == kaggle_dataset.shape[0]
assert kaggle_subset_C.shape[1] == kaggle_dataset.shape[1]
assert sum(kaggle_subset_C.CurrentJobTitleSelect == 'Programmer') == 0

```

<!-- grade_id: cell-c928e9856866109b, points: 0, grade: false, solution: true -->
```python
def kaggle_dataset_D(df):
    """ 
    Show only Computer Scientists in dataframe
    
    Args:
        df (pd.DataFrame): the input DataFrame

    Returns:
        (pd.DataFrame): output DataFrame

    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fa5afc2398e754bf, points: 2, grade: true, solution: false -->
```python
# This cell is what will test your code, please ignore it!

kaggle_dataset = pd.read_csv('data/multiple_choice_2017.csv')
kaggle_subset_D = kaggle_dataset_D(kaggle_dataset)

assert isinstance(kaggle_subset_D, pd.DataFrame)
assert kaggle_subset_D.shape[0] == kaggle_dataset.shape[0]
assert kaggle_subset_D.shape[1] == kaggle_dataset.shape[1]
assert sum(kaggle_subset_D.dropna().CurrentJobTitleSelect != 'Computer Scientist') == 0

```

Now look into the clues and retrieve the correct values to fill the puzzle:

* **Clue 9 - The first word of the employment status of the 4th element of subset D**
* **Clue 10 - The top compensation currency on subset C**


```python
NINTH_CLUE = kaggle_subset_D.EmploymentStatus.values[4].split(' ')[0]
TENTH_CLUE = kaggle_subset_C.CompensationCurrency.value_counts().index[0] 

```

Now that you have all the clues, check the puzzle for the secret key

```python
draw_final_puzzle(
    [
        FIRST_CLUE, 
        SECOND_CLUE, 
        THIRD_CLUE, 
        FORTH_CLUE, 
        FIFTH_CLUE, 
        SIXTH_CLUE, 
        SEVENTH_CLUE, 
        EIGTH_CLUE, 
        NINTH_CLUE, 
        TENTH_CLUE
    ]
)
```

<!-- grade_id: cell-029a11bdce1cfb7c, points: 0, grade: false, solution: true -->
```python
# Introduce the highlighted words you see, in the following form:
# kaggle_key = "highlightedword"

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2fd6aa23728ab5d6, points: 0, grade: true, solution: false -->
```python
assert hashlib.sha256(kaggle_key.encode()).hexdigest() == "910a9c5274ba0637ca5882fdef4190e608fb05e465da46518bd7f2fe2eb6d93d"

```

Congratulations, you made it! You would now be able to enter the actual challenge and brag to all your friends about how good you are in data science 😄 

<img src="images/excel.jpg" alt="excel" width="40%"/>


