# BLU05  - Exercise Notebook 

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
from sklearn.metrics import mean_absolute_error, mean_squared_error
from statsmodels.tsa.seasonal import seasonal_decompose
from statsmodels.graphics.tsaplots import plot_acf, plot_pacf
import hashlib # for grading purposes

import pmdarima as pm
from pmdarima.pipeline import Pipeline
from pmdarima.preprocessing import BoxCoxEndogTransformer
import warnings
warnings.simplefilter(action='ignore', category=FutureWarning)   
warnings.filterwarnings("ignore") # specify to ignore warning messages
%matplotlib inline
plt.rcParams['figure.figsize'] = (12, 4)
```

# Functions

```python
def load_emissions_data():
    emissions = pd.read_csv('data/emissions.csv')
    emissions = emissions[emissions.Description=='Coal Electric Power Sector CO2 Emissions']
    emissions['YYYYMM'] = emissions['YYYYMM'].astype(str)
    emissions['YYYYMM'] = pd.to_datetime(emissions['YYYYMM'],format='%Y%m', errors='coerce').dropna()
    emissions = emissions.set_index('YYYYMM')
    emissions = emissions.sort_index()
    emissions['Value'] = emissions.Value.astype(float)
    emissions.index = emissions.index.rename('date')
    emissions = emissions['Value']
    emissions = emissions['1980':'2000']
    emissions = emissions.resample('MS').mean()
    return emissions
```

```python
def load_coal_data():
    np.random.seed(10)
    df = pd.read_csv('data/MER_T06_01.csv')
    df = df[df.Description=='Coal Consumption']
    df.YYYYMM = pd.to_datetime(df.YYYYMM, format='%Y%m',errors='coerce' )
    df = df.set_index('YYYYMM')
    df = df.loc['1980':'2000']
    df.Value = pd.to_numeric(df.Value)
    df.Value = df.Value + np.random.normal(0,df.Value.mean()/20,len(df.Value))
    return df.Value.to_frame()
```

# Let's predict CO2 emissions!

```python
emissions = load_emissions_data()
```

```python
emissions_train = emissions[:'1997']
emissions_test = emissions['1997':]
```

##### Let's get an overview of the df and plot the series to get an idea of what's going on

```python
emissions_train.head(13)
```

```python
emissions_train.plot();
plt.xlabel('Time');
plt.title('Monthly CO2 emissions (gCO2eq)');
```

### Q1:Getting a feel of the data

<!-- grade_id: cell-20ec8c1b7ca5a846, points: 0, grade: false, solution: true -->
```python
# Is the magnitude of the variance changing? answer with 'Yes' or 'No'
# variance_change = 

# Does it have an apparent trend? answer with 'Yes' or 'No'
# apparent_trend = 

# the emissions plot can give you some insights on these points.
# but we've also talked about clearer approaches like using some rolling ???? function
# and the seasonal_decompose function (extra tip: you can use a 1-year window, window=12)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d1e64704448ac7a8, points: 0.5, grade: true, solution: false -->
```python
expected_hash = '85a39ab345d672ff8ca9b9c6876f3adcacf45ee7c1e2dbd2408fd338bd55e07e'
assert hashlib.sha256(variance_change.encode()).hexdigest() == expected_hash
expected_hash = '85a39ab345d672ff8ca9b9c6876f3adcacf45ee7c1e2dbd2408fd338bd55e07e'
assert hashlib.sha256(apparent_trend.encode()).hexdigest() == expected_hash
```

### Q2: Make the time series stationary

<!-- grade_id: cell-4673648012938702, points: 0, grade: false, solution: true -->
```python
# emissions_train_log = 
# emissions_train_log_detrend = 
# emissions_train_log_detrend_without_nans =  

# YOUR CODE HERE
raise NotImplementedError()
```

```python
plot_acf(emissions_train);
plt.xlabel('Lag')
plot_acf(emissions_train_log);
plt.xlabel('Lag')
plot_acf(emissions_train_log_detrend_without_nans);
plt.xlabel('Lag')
```

<!-- grade_id: cell-37f19bfc31b2d1e0, points: 2, grade: true, solution: false -->
```python
expected_hash = 'a9bd2829b2d2191c657f4e3b13bf0feca647eb3630a15ebac45aa578d6199e3d'
assert hashlib.sha256(str(np.round(emissions_train_log_detrend_without_nans[0],5)).encode()).hexdigest() == expected_hash
```

### Optional: Visualize the variance and/or trend of the stationarity timeseries

### Q3: Looking for seasonality

<!-- grade_id: cell-b2bfb6431053bbb1, points: 0, grade: false, solution: true -->
```python
# Use the autocorrelation plots to infer what is the seasonality of the time series.

#plot_pacf(what goes here?), if you're curious take also a look at the acf
#plt.xlabel('lag')
#plt.ylabel('Autocorrelation')
#plt.show()

# S = 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-31fe58e5b06199fb, points: 1.5, grade: true, solution: false -->
```python
assert isinstance(S, int)
expected_hash = '6b51d431df5d7f141cbececcf79edf3dd861c3b4069f0b11661a3eefacbba918'
assert hashlib.sha256(str(S).encode()).hexdigest() == expected_hash
```

### Q4: Fit a SARIMAX model on the train set using AutoARIMA

```python
S
```

<!-- grade_id: cell-6312e30d5d2437ff, points: 0, grade: false, solution: true -->
```python
# Use S according to what you saw on the previous question.
# Use the logged train dataset, but not the diffed since pmdarima can find that on his own.
# Use the 'nm' method and maxiter=20
# Also define that warnings should be suppressed

# sarimax =
# sarimax.fit()

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-70c443f3d02f4229, points: 2, grade: true, solution: false -->
```python
expected_hash = '5eed90f880c272ad2ae830b7d6c54f8362876cebc88c946db10dbf26aafa9a1b'
assert hashlib.sha256(str(sarimax.get_params()).encode()).hexdigest() == expected_hash
```

### Q5: Use that model to do the in-sample predictions and calculate the corresponding MSE

<!-- grade_id: cell-d322536e95c10f9b, points: 0, grade: false, solution: true -->
```python
# Don't forget that you used the logged dataset when training the model, thus its outputs will also be logged

# predictions =
# mse = 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b8b5a48f82f2b6a3, points: 2, grade: true, solution: false -->
```python
expected_hash = 'e84be574290605e678015094ff2db55a865d20754df25ef1fbbd86b9cc73a576'
assert hashlib.sha256(str(np.round(predictions,1)).encode()).hexdigest() == expected_hash

expected_hash = 'd4d63176def902021ddfc53b452cc750fcbec8e3e087976502ecf529ecaa2b8e'
assert hashlib.sha256(str(np.round(mse,1)).encode()).hexdigest() == expected_hash
```

### Q6: Forecast the test set using a multi-step forecast and calculate the corresponding MSE

<!-- grade_id: cell-893836dce142cc23, points: 0, grade: false, solution: true -->
```python
# Don't forget that you used the logged dataset when training the model.

# forecast = 
# mse = 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a9c4b093d376d0be, points: 2, grade: true, solution: false -->
```python
expected_hash = '433aed5fef7646ce383023e7fe95ebc0f4291a215cbb3d61fe3f3d307c06d3b4'
assert hashlib.sha256(str(np.round(forecast,1)).encode()).hexdigest() == expected_hash

expected_hash = '2185f687cc1bbae4025d1b867cd762892de64fac607059896b6bce8c62f62f34'
assert hashlib.sha256(str(np.round(mse,1)).encode()).hexdigest() == expected_hash
```

```python
emissions_test.plot()
pd.Series(forecast, index=emissions_test.index).plot()
plt.xlabel('Time')
plt.title('Monthly CO2 emissions (gCO2eq)')
plt.legend(['original','forecast'],loc=2)
print(mse)
```

### Q7: Forecast the test set using multiple one-step forecasts

### Q7.1: Get the rolling one-step forecasts

<!-- grade_id: cell-7df87d54c98920a1, points: 0, grade: false, solution: true -->
```python
# Don't forget to fit your sarimax model again
# You will need to make a for loop that updates the model

# sarimax_forecast_one_step_refit = []
# sarimax.fit(emissions_train_log)  
# emissions_test_logged = 
# Update loop

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bb5f9d7036959ca8, points: 4, grade: true, solution: false -->
```python
expected_hash = '7080152205ba48d591d1dede83e16a46e27f6d228a709220cb4e31d80a46adcc'
assert hashlib.sha256(str(np.round(sarimax_forecast_one_step_refit)).encode()).hexdigest() == expected_hash
```

### Q7.2: Calculate the corresponding MSE

<!-- grade_id: cell-b08743bd56aa6f39, points: 0, grade: false, solution: true -->
```python
# Don't forget that you used the logged dataset when training the model.

# sarimax_forecast_one_step_refit_exp =
# mse =

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d1d8e5e079f2aeab, points: 1, grade: true, solution: false -->
```python
expected_hash = 'bba8868023dbbbb0b719df3b375500c3a9643db09f8919a43a7774d578adc329'
assert hashlib.sha256(str(np.round(sarimax_forecast_one_step_refit_exp,1)).encode()).hexdigest() == expected_hash

expected_hash = 'cb30211d8e99f76a7f0b4f080b48cb2095f6b6128b07d52bbf030a78c146aea0'
assert hashlib.sha256(str(np.round(mse,1)).encode()).hexdigest() == expected_hash
```

```python
emissions_test.plot()
pd.Series(sarimax_forecast_one_step_refit_exp, index=emissions_test.index).plot()
plt.xlabel('Time')
plt.title('Monthly CO2 emissions (gCO2eq)')
plt.legend(['original','forecast'],loc=2)

print(mse)
```

### Q8: Repeat Q7 now with an exogenous input

Let's test the performance of our model using an exogenous input: the coal consumption. Note that the dataset below contains the coal consumption forecast made on the previous month. This means that we can use this exogenous data for doing one-step forecasts. 

```python
exog = load_coal_data()
exog_train = exog[:'1997']
exog_test = exog['1997':]
```

Let's look at the data to get an idea of what we can expect:

```python
exog_train.plot();
plt.xlabel('Time');
plt.title('Monthly coal consumption');
```

### Q8.1: Get the rolling one-step forecasts

<!-- grade_id: cell-31d0dfc073a764a7, points: 0, grade: false, solution: true -->
```python
# Don't forget to fit your sarimax model again
# You will need to make a for loop that updates the model
# Note that this case is slightly different from the learning notebooks, where we used the exogenous data to do a
# multi-step forecast. Here we need to use the exogenous data in the context of multiple one-step forecasts, this 
# means we need to have an update loop where we also include the new values of the exogenous data 
# (similar to what was done in Q7.1).

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a89774a4f207e248, points: 4, grade: true, solution: false -->
```python
expected_hash = 'acb1cf9da851c0c91595a3c6f68cc4669dbcb9d23dde524100b656376e1de214'
assert hashlib.sha256(str(np.round(sarimax_forecast_one_step_refit,1)).encode()).hexdigest() == expected_hash
```

### Q8.2: Calculate the corresponding MSE

<!-- grade_id: cell-08cf8f9cc63a4753, points: 0, grade: false, solution: true -->
```python
# Don't forget that you used the logged dataset when training the model.

# sarimax_forecast_one_step_refit_exp = 
# mse = 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-44a5908d26970bf6, points: 1, grade: true, solution: false -->
```python
expected_hash = 'ba39644cf33119587851eb252e85132508761fd8a8d708ca2e8e314754d8b9d9'
assert hashlib.sha256(str(np.round(sarimax_forecast_one_step_refit_exp,1)).encode()).hexdigest() == expected_hash

expected_hash = '0ac326686f62e06176321c58953c4eb0fed6c7610bdef09b71c7e8c36424daac'
assert hashlib.sha256(str(np.round(mse,1)).encode()).hexdigest() == expected_hash
```

```python
emissions_test.plot()
pd.Series(sarimax_forecast_one_step_refit_exp, index=emissions_test.index).plot()
plt.xlabel('Time')
plt.title('Monthly CO2 emissions (gCO2eq)')
plt.legend(['original','forecast'],loc=2)
print(mse)
```

### Q666: Unlock the power of the WMSE (optional)
So here's a little something which we didn't discuss in the Learning materials: _weighted_ metrics (be it MAE, MSE, etc.)

Now, let's dive in a trip of massive imagination:
- the CO2 emissions data set is respective to the _Island of Wonders_ 
- you are a travel agent selling holiday packs for this island (say what?!)
- every month you sell the same number of packs, except for August where your sales double (summer vacation!)
- all your clients are very picky about urban pollution and ask you for a month-ahead CO2 forecast

If you train/evaluate your model as usual, you get an all-rounder tool and an overall assessment. However, it might be of interest to take into account that there are 2x the clients in August and maybe give this month 2x the relevance.

For simplicity, we will leave the "training the model part" out (you can still explore this on your own). But take a look at the `sample_weight` parameter of `mean_squared_error` ([link](https://scikit-learn.org/stable/modules/generated/sklearn.metrics.mean_squared_error.html)) and try to **calculate the test MSE using the same `sarimax_forecast_one_step_refit_exp` forecast but giving 2x the importance to the August records**.



<!-- grade_id: cell-c48beed0f9e0c069, points: 0, grade: false, solution: true -->
```python
# create a baseline weight list filled with ones (1)
# w = 

# create a list with the month for each element of the test set
# month = 

# edit the w list, giving the August elements double the weight
# ...

# calculate the WMAE
# wmse_august_double_relevance = 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f420dc7a8b6051bc, points: 0, grade: true, solution: false -->
```python
expected_hash = '0ac326686f62e06176321c58953c4eb0fed6c7610bdef09b71c7e8c36424daac'
assert hashlib.sha256(str(np.round(mse,1)).encode()).hexdigest() == expected_hash
```

```python
print(mse)
print(wmse_august_double_relevance)
```