# SLU03 | Visualization with Pandas & Matplotlib: Exercise notebook
***

## Start by importing these packages

```python
import pandas as pd
import numpy as np

# for evaluation purposes
import hashlib
from plotchecker import PlotChecker

def _hash(s):
    return hashlib.blake2b(
        bytes(str(s), encoding='utf8'),
        digest_size=5
    ).hexdigest()

def get_data(p, ax=0):
    all_x_data = []
    lines = p.axis.get_lines()
    collections = axis.collections
    if len(lines) > 0:
        all_x_data.append(np.concatenate([x.get_xydata()[:, ax] for x in lines]))
    if len(collections) > 0:
        all_x_data.append(np.concatenate([x.get_offsets()[:, ax] for x in collections]))
    return np.concatenate(all_x_data, axis=0)
```

In this notebook you will practice the following:

- Basics of Matplotlib
- Scatter plots
- Line charts
- Bar/column charts
- Histograms
- Box plots

To learn about data visualization, we are going to use a **modified version** of [Zomato Bangalore Restaurants](https://www.kaggle.com/himanshupoddar/zomato-bangalore-restaurants) which has information about restaurants

The dataset is located at `data/zomato.csv`, and has the following fields:


>**`name`**: Restaurant name.  
>**`vote_rate`**: Average rating in zomato.  
>**`vote_count`**: Number of votes in zomato.  
>**`location`**: Location where the restaurant is based.  
>**`type`**: Restaurant type.  
>**`cuisine`**: Restaurant main dish cuisine type.  
>**`cost_2pax`**: Approximate meal-cost for 2 people (in $).  


```python
restaurants = pd.read_csv("data/zomato.csv")
```

```python
restaurants.shape
```

```python
restaurants.head()
```

***
## Exercise 1

First, import matplotlib, pyplot and the matplotlib inline ✨magic✨. 

<!-- grade_id: cell-96c602a1da8f81e5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f150c73eacb8515f, points: 1, grade: true, solution: false -->
```python
assert plt, "Did you import module pyplot from matplotlib and called it plt?"

print('--- Success ---')
```

And, change the default chart size to 12 inches width and 12 inches height

<!-- grade_id: rcParams_answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: rcParams_test, points: 1, grade: true, solution: false -->
```python
s = [plt.rcParams["figure.figsize"][0], plt.rcParams["figure.figsize"][1]]
s.sort()

assert _hash(s) in ['9b75178bbe', 'e139d197c5'], "Did you change the default figure sizes?"

print('--- Success ---')
```

<hr>

## 🛑 Important note about the grading

**Grading plots is difficult**, we are using `plotchecker` to grade the plots with `nbgrader`. 
For `plotchecker` to work with `nbgrader`, **we need to add on each cell**, the line

> **`axis = plt.gca();`**

**After the code required** to do the plot.

<div class="alert alert-danger">
    ⚠️ <b>NOTE: </b>If you get the <b>ImportError</b>, <i>plotchecker not defined</i>, make sure you activate the right environment for this unit!
</div>

> **For example**, *if we want to plot a `scatter plot` showing the relationship between **`cost for 2 people`** and **`vote rating`** among the restaurants we would do as follows:*

```python
# code required to plot
restaurants.plot.scatter(x="cost_2pax", y="vote_rate")

# last line in the cell required to "capture" the cell and being able to grade it with nbgrader
axis = plt.gca();
```

***

## Exercise 2

Let's see how does the **`cost for 2 people`** correlate with the **`number of votes`**?

<!-- grade_id: cell-d1abe51868ed32d0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
axis = plt.gca();
```

<!-- grade_id: cell-dd0d5562ad8f2d75, points: 3, grade: true, solution: false -->
```python
pc = PlotChecker(axis)
data = get_data(pc)
l = [pc.xlabel] + [pc.ylabel]
l.sort()

assert _hash(len(data)) == '4b7eb88930', "Did you set the right variables for the plot axes?"
assert _hash(l) in ['d1ed203167'], "Did you set the right variables for the plot axes?"
assert _hash(get_data(pc,1)) in ['676a41f8f7', '68f9b6d00d'], "Did you set the right variables for the plot axes?"

print('--- Success ---')
```

## Exercise 3

How the average **`cost for 2 people`** are trending through the **`rating votes`**?

- set the plot **title** to **`Average Restaurant cost by Rating`**.
- label the **x-axis** as **`Zomato Ratings`**.
- label the **y-axis** as **`Restaurant Cost (2pax)`**.

<div class="alert alert-info">
    To calculate the average cost by restaurant rating we need to perform an <b><a href='https://jakevdp.github.io/PythonDataScienceHandbook/03.08-aggregation-and-grouping.html'>aggregation</a></b>, pandas support this by doing a technique called <b><a href='(https://pandas.pydata.org/pandas-docs/stable/user_guide/groupby.html'>Split-Apply-Combine</a></b>. <b>This will be explained in the Data Wrangling Specialization.</b> For now we will do the grouping for you:
</div>

```python
avg_cost_2pax_by_rating = restaurants.groupby("vote_rate")["cost_2pax"].mean().reset_index()
avg_cost_2pax_by_rating.columns = ["rating", "avg_cost_2pax"]
avg_cost_2pax_by_rating.head()
```

<div class="alert alert-warning">
⚠️ <b>NOTE: </b>Make sure you use the dataframe named <b>avg_cost_2pax_by_rating</b> for the exercise
</div>

<!-- grade_id: cell-73b4bfa635b491d8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
axis = plt.gca();
```

<!-- grade_id: cell-e810ccdca9b86e40, points: 3, grade: true, solution: false -->
```python
pc = PlotChecker(axis)
l = [pc.xlabel] + [pc.ylabel]
l.sort()

assert _hash(sorted(get_data(pc))) in ['bc499f7725'], "Did you set the right variables for the plot axes?"
assert _hash(sorted(get_data(pc,1))) in ['1062114291'], "Did you set the right variables for the plot axes?"
assert _hash(pc.title) in ['b9c481c4ef'], "Did you set the right plot title?"
assert _hash(l) in ['640d7ee8aa'], "Did you set the right plot axes?"

print('--- Success ---')
```

## Exercise 4

How does the median **`number of votes`** vary by restaurant **`type`**? 

- label the **x-axis** as **`Number of Votes`**.
- label the **y-axis** as **`Restaurant Type`**.
- change the **plot color** to **`chocolate`**.

Again, we will do the **grouping** for you:

```python
median_votes_by_type = restaurants.groupby("type")["vote_count"].median().reset_index()
median_votes_by_type.columns = ["type", "median_votes"]
median_votes_by_type
```

<div class="alert alert-warning">
⚠️ <b>NOTE: </b>Make sure you use the dataframe named <b>median_votes_by_type</b> for the exercise
</div>

<!-- grade_id: cell-f7836531abe09da4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
axis = plt.gca();
```

<!-- grade_id: cell-fcd988c35854556f, points: 3, grade: true, solution: false -->
```python
pc = PlotChecker(axis)
pc._patches = np.array(pc.axis.patches)
pc._patches = pc._patches[np.argsort([p.get_x() for p in pc._patches])]

pc.widths = np.array([p.get_width() for p in pc._patches])
pc.heights = np.array([p.get_height() for p in pc._patches])

l = [pc.xlabel] + [pc.ylabel]
l.sort()

assert _hash(len(pc._patches)) in ['1cb85c3d73'], "Did you set the right categorical variable and the right plot type?"
assert _hash(np.sort(pc.widths)) in ['8f9f8c51a9'], "Did you set the right categorical variable and the right plot type?"
assert _hash(l) in ['bcd1693b58'], "Did you set the right plot axes?"
assert _hash([pc._patches[i].get_facecolor() for i,n in enumerate(pc._patches)]) == '1fe2bb456d', "Did you change the plot color?"

print('--- Success ---')
```

## Exercise 5

How does the average **`number of votes`** compares by the **`rating`**? 

- label the **x-axis** as **`Zomato rating`** and **size** **`20`**.
- label the **y-axis** as **`Number of Votes`** and **size** **`20`**.
- set the plot **title** to **`Average number of votes comparison by Rating`** and **size** **`25`**.

Doing the **grouping** for you:

```python
avg_vote_by_rating = restaurants.groupby("vote_rate")["vote_count"].mean().reset_index()
avg_vote_by_rating.columns = ["rating", "avg_votes"]
avg_vote_by_rating
```

<div class="alert alert-warning">
⚠️ <b>NOTE: </b>Make sure you use the dataframe named <b>avg_vote_by_rating</b> for the exercise
</div>

<!-- grade_id: cell-372fe0ddfdb5ca15, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
axis = plt.gca();
```

<!-- grade_id: cell-aaf06b8c5ace2004, points: 3, grade: true, solution: false -->
```python
pc = PlotChecker(axis)
pc._patches = np.array(pc.axis.patches)
pc._patches = pc._patches[np.argsort([p.get_x() for p in pc._patches])]

pc.widths = np.array([p.get_width() for p in pc._patches])
pc.heights = np.array([p.get_height() for p in pc._patches])

l = [pc.xlabel] + [pc.ylabel]
l.sort()

assert _hash(len(pc._patches)) in ['d15bce6d7d'], "Did you set the right variable and the right plot type?"
assert _hash(np.sort(pc.widths)) in ['1050369a36'], "Did you set the right variable and the right plot type?"
assert _hash(np.sort(pc.heights)) in ['51713c2dfb'], "Did you set the right variable and the right plot type?"
assert _hash(pc.title) in ['3db42086ff'], "Did you set the right plot title?"
assert _hash(l) in ['ce43ba93aa'], "Did you set the right plot axes?"

print('--- Success ---')
```

## Exercise 6

How is the variable **`vote_rate`** distributed through the better ratings? 
- show only restaurants with a minimum rating of **4 points** (up to 5 points). 
> 💡 **hint:** _set a limit_ 
- change the number of **bins to 30**.
- change the **plot color** to **`yellowgreen`**.

<!-- grade_id: cell-1988d757a6ed50ea, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
axis = plt.gca();
```

<!-- grade_id: cell-da6b717de847354d, points: 3, grade: true, solution: false -->
```python
pc = PlotChecker(axis)
pc._patches = np.array(pc.axis.patches)
pc._patches = pc._patches[np.argsort([p.get_x() for p in pc._patches])]
pc.widths = np.array([p.get_width() for p in pc._patches])
pc.heights = np.array([p.get_height() for p in pc._patches])

assert _hash(np.sort(pc.widths)) in ['1717c26bdc'], "Did you set the right plot type? You should not change the dataset."
assert _hash(np.sort(pc.heights)) in ['93b03e2bf6'], "Did you set the right plot type? You should not change the dataset."
assert _hash(pc.xlim[0]) == '8599ebffc2', "Did you set the right limit?"
assert _hash(pc.xlim[1]) == 'f24d02e0e6', "Did you set the right limit?"
assert _hash([pc._patches[i].get_facecolor() for i,n in enumerate(pc._patches)]) == '697422ff42', "Did you change the plot color?"

print('--- Success ---')
```

## Exercise  7

Now, we will analyze the restaurants of Chinese food and compare them between the rating and the cuisine type of those restaurants. Following the steps below:

- filter the **`restaurants`** dataset by **chinese** **`cuisine`**.
- save them in a new **dataframe** called **`chinese_restaurants`**.
- make a plot that displays the **rating** by chinese restaurant **type** and that allows us to check if there are outliers.
- change the default **plot style** to **`seaborn`**.
- set the plot argument **vert** to **`False`**.
- set the plot argument **figsize** to **`10`** inches width and **`10`** inches height.
- set the plot argument **grid** to **`False`**
- set the plot argument **fontsize** to **`14`**

<!-- grade_id: cell-da1cf313df52ced6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
axis = plt.gca();
```

<!-- grade_id: cell-d67c3ae944e4adb7, points: 3, grade: true, solution: false -->
```python
pc = PlotChecker(axis)
pc._lines = pc.axis.get_lines()
pc.colors = np.array([pc._color2rgb(x.get_color()) for x in pc._lines])

assert _hash(chinese_restaurants['cuisine'].unique()) in ['3c110a6b03'], "Did you filter the chinese restaurtants?"
assert _hash(pc.colors[0]) in ['e68950262e'], "Did you set the right plot style?"
assert _hash(pc.yticklabels) in ['474e6af7df'], "Did you select the right categorical variable for the plot?"
assert _hash(pc.xticklabels) in ['f50b6dfe86'], "Did you select the right categorical variable for the plot?"

print('--- Success ---')
```


# 🏁 Ungraded Exercise 🏁
Load the file misterious_data.csv and use data visualization to answer the following questions:

* How is the distribution of **x** in general?
* Are there any **outlier** in any of the fields?
* Which **2 charts** better represent the underlying data?
* Change their **style** to **`bmh`**.
* Add **titles** to each chart explaining them.
